Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Dmitry Chudinovskikh"
COLLABORATORS = "None"
```

---

# Exercise 1 - Shell basics

Work through as much of the Software Carpentry [lesson on the Unix Shell](http://swcarpentry.github.io/shell-novice/) as you can.  Run through the Setup section just below, then open a terminal through Jupyter to run through the exercises.

After you have completed the first few sections of the tutorial, return to this notebook.  Execute all of the cells, and answer all of the questions.


## Setup - getting required files 

To get started, you'll need to have the required files in your directory.  Use `wget` to get them:

```python
!wget http://swcarpentry.github.io/shell-novice/data/shell-novice-data.zip
```

Output:
```
--2016-09-01 14:27:48--  http://swcarpentry.github.io/shell-novice/data/shell-novice-data.zip
Resolving swcarpentry.github.io (swcarpentry.github.io)... 151.101.32.133
Connecting to swcarpentry.github.io (swcarpentry.github.io)|151.101.32.133|:80... connected.
HTTP request sent, awaiting response... 200 OK
Length: 168355 (164K) [application/zip]
Saving to: ‘shell-novice-data.zip.1’

shell-novice-data.z 100%[===================>] 164.41K  --.-KB/s    in 0.001s  

2016-09-01 14:27:49 (165 MB/s) - ‘shell-novice-data.zip.1’ saved [168355/168355]
```

```python
!unzip -l shell-novice-data.zip
```

Output:
```
Archive:  shell-novice-data.zip
  Length      Date    Time    Name
---------  ---------- -----   ----
        0  2015-11-22 16:39   data-shell/
      199  2015-11-22 16:39   data-shell/.bash_profile
        0  2015-11-17 07:02   data-shell/creatures/
     1838  2015-11-17 07:02   data-shell/creatures/basilisk.dat
     1833  2015-11-17 07:02   data-shell/creatures/unicorn.dat
        0  2015-11-17 07:02   data-shell/data/
      283  2015-11-17 07:02   data-shell/data/amino-acids.txt
      136  2015-11-17 07:02   data-shell/data/animals.txt
        0  2015-11-17 07:02   data-shell/data/elements/
      270  2015-11-17 07:02   data-shell/data/elements/Ac.xml
      312  2015-11-17 07:02   data-shell/data/elements/Ag.xml
      314  2015-11-17 07:02   data-shell/data/elements/Al.xml
      325  2015-11-17 07:02   data-shell/data/elements/Am.xml
      307  2015-11-17 07:02   data-shell/data/elements/Ar.xml
      326  2015-11-17 07:02   data-shell/data/elements/As.xml
      261
```

```python
!unzip shell-novice-data.zip
```

Output:
```
Archive:  shell-novice-data.zip
   creating: data-shell/
  inflating: data-shell/.bash_profile  
   creating: data-shell/creatures/
  inflating: data-shell/creatures/basilisk.dat  
  inflating: data-shell/creatures/unicorn.dat  
   creating: data-shell/data/
  inflating: data-shell/data/amino-acids.txt  
  inflating: data-shell/data/animals.txt  
   creating: data-shell/data/elements/
  inflating: data-shell/data/elements/Ac.xml  
  inflating: data-shell/data/elements/Ag.xml  
  inflating: data-shell/data/elements/Al.xml  
  inflating: data-shell/data/elements/Am.xml  
  inflating: data-shell/data/elements/Ar.xml  
  inflating: data-shell/data/elements/As.xml  
  inflating: data-shell/data/elements/At.xml  
  inflating: data-shell/data/elements/Au.xml  
  inflating: data-shell/data/elements/B.xml  
  inflating: data-shell/data/elements/Ba.xml  
  inflating: data-shell/data/elements/Be.xml  
  inflating: data-shell/data/elements/Bi.xml  
  inflating: data-shell/dat
```

*Note*: you only need to do this once per session while using Jupyter on datanotebook.org.  You can open a terminal now and work through the steps, and return to this notebook a little later, and the files will be available either way.  That's because you're working in the same local directory.

**However**!  If you download this file, close your Jupyter session for the night, then come back tomorrow and open up a new Jupyter session on the server again, you'll need to get those sample files again.  Just execute the cells above to do it.

Okay, let's get on with the exercise!

## Navigating Files and Directories

As you work through this section of the tutorial, complete the steps here as well, using the `!` shell escape command.  Execute each cell as you go.

These steps aren't exactly the same as what's in the tutorial, where the file layout is a little different and where they're not using a notebook like we are.  That's okay.  Just consider this practice.

```python
!whoami
```

Output:
```
dmitrychu
```

```python
!pwd
```

Output:
```
/home/dmitrychu
```

```python
!ls -F
```

Output:
```
anaconda3/			 Music/			  Templates/
Anaconda3-4.1.1-Linux-x86_64.sh  numbers.txt		  Untitled1.ipynb
data-shell/			 Pictures/		  Untitled Folder/
Desktop/			 Public/		  Untitled.ipynb
Documents/			 R/			  Videos/
Downloads/			 shell-novice-data.zip	  wk2.ipynb
examples.desktop		 shell-novice-data.zip.1
exercise-01.ipynb		 temp/
```

```python
!ls -F
```

Output:
```
anaconda3/			 Music/			  Templates/
Anaconda3-4.1.1-Linux-x86_64.sh  numbers.txt		  Untitled1.ipynb
data-shell/			 Pictures/		  Untitled Folder/
Desktop/			 Public/		  Untitled.ipynb
Documents/			 R/			  Videos/
Downloads/			 shell-novice-data.zip	  wk2.ipynb
examples.desktop		 shell-novice-data.zip.1
exercise-01.ipynb		 temp/
```

```python
!ls -F data-shell/
```

Output:
```
creatures/  Desktop/	north-pacific-gyre/  pizza.cfg	writing/
data/	    molecules/	notes.txt	     solar.pdf
```

```python
!ls -aF
```

Output:
```
./				 exercise-01.ipynb    .Rhistory
../				 .gconf/	      .rstudio-desktop/
anaconda3/			 .gnupg/	      shell-novice-data.zip
Anaconda3-4.1.1-Linux-x86_64.sh  .ICEauthority	      shell-novice-data.zip.1
.atom/				 .ipynb_checkpoints/  .spyder2/
.bash_history			 .ipython/	      .spyder2-py3/
.bash_logout			 .jupyter/	      .sudo_as_admin_successful
.bashrc				 .local/	      temp/
.cache/				 .mozilla/	      Templates/
.config/			 Music/		      Untitled1.ipynb
.continuum/			 numbers.txt	      Untitled Folder/
data-shell/			 Pictures/	      Untitled.ipynb
Desktop/			 .pki/		      Videos/
.dmrc				 .profile	      wk2.ipynb
Documents/			 Public/	      .Xauthority
Downloads/			 .python_history      .xsession-errors
examples.desktop		 R/		      .xsession-errors.old
```

```python
!ls -af .
```

Output:
```
.rstudio-desktop	   .python_history     Public
.config			   .continuum	       Anaconda3-4.1.1-Linux-x86_64.sh
.xsession-errors.old	   .local	       ..
Pictures		   .bash_logout        .spyder2
.ipython		   .xsession-errors    exercise-01.ipynb
Downloads		   .atom	       Music
Untitled1.ipynb		   anaconda3	       data-shell
shell-novice-data.zip	   .		       .jupyter
Templates		   .gnupg	       .Xauthority
numbers.txt		   R		       Untitled Folder
Videos			   .profile	       .spyder2-py3
.Rhistory		   .bashrc	       .bash_history
.sudo_as_admin_successful  .mozilla	       .cache
temp			   .gconf	       Desktop
examples.desktop	   .ICEauthority       .pki
Untitled.ipynb		   .ipynb_checkpoints  .dmrc
shell-novice-data.zip.1    wk2.ipynb	       Documents
```

What is the difference between the two previous cells, and what does the single dot mean?

<!-- grade_id: 156, points: 1, grade: true, solution: true -->
In the first cell all the entries are sorted and a trailing is added to the names of directories, whereas in the second cell all the entries are not sorted and listed in directory order. Single dot stands for current directory

```python
!ls -F ..
```

Output:
```
dmitrychu/
```

What do the double dots mean?

<!-- grade_id: 155, points: 1, grade: true, solution: true -->
Double dots make the system go up to the directory one level above the one we are currently working with. 

```python
!ls data-shell/north-pacific-gyre/2012-07-03/
```

Output:
```
goodiff		NENE01736A.txt	NENE01843A.txt	NENE01978B.txt	NENE02040Z.txt
goostats	NENE01751A.txt	NENE01843B.txt	NENE02018B.txt	NENE02043A.txt
NENE01729A.txt	NENE01751B.txt	NENE01971Z.txt	NENE02040A.txt	NENE02043B.txt
NENE01729B.txt	NENE01812A.txt	NENE01978A.txt	NENE02040B.txt
```

## Working with Files and Directories

The following cells come from the next section of the tutorial.

```python
!ls -F
```

Output:
```
anaconda3/			 Music/			  Templates/
Anaconda3-4.1.1-Linux-x86_64.sh  numbers.txt		  Untitled1.ipynb
data-shell/			 Pictures/		  Untitled Folder/
Desktop/			 Public/		  Untitled.ipynb
Documents/			 R/			  Videos/
Downloads/			 shell-novice-data.zip	  wk2.ipynb
examples.desktop		 shell-novice-data.zip.1
exercise-01.ipynb		 temp/
```

```python
!mkdir thesis
```

<!-- grade_id: 121, points: 1, grade: true, solution: false -->
```python
import os
assert "thesis" in os.listdir()
```

```python
!ls -F
```

Output:
```
anaconda3/			 Music/			  Templates/
Anaconda3-4.1.1-Linux-x86_64.sh  numbers.txt		  thesis/
data-shell/			 Pictures/		  Untitled1.ipynb
Desktop/			 Public/		  Untitled Folder/
Documents/			 R/			  Untitled.ipynb
Downloads/			 shell-novice-data.zip	  Videos/
examples.desktop		 shell-novice-data.zip.1  wk2.ipynb
exercise-01.ipynb		 temp/
```

You can't use the nano editor here in Jupyter, so we'll use the `touch` command to create an empty file instead.

```python
!touch thesis/draft.txt
```

<!-- grade_id: 123, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir("thesis")
```

```python
!ls -F thesis
```

Output:
```
draft.txt
```

Removing files and directories.

```python
!rm thesis/draft.txt
```

<!-- grade_id: 125, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" not in os.listdir("thesis")
```

```python
!rm thesis
```

Output:
```
rm: cannot remove 'thesis': Is a directory
```

```python
!rmdir thesis
```

<!-- grade_id: 127, points: 1, grade: true, solution: false -->
```python
assert "thesis" not in os.listdir()
```

```python
!ls
```

Output:
```
anaconda3			 Music			  Templates
Anaconda3-4.1.1-Linux-x86_64.sh  numbers.txt		  Untitled1.ipynb
data-shell			 Pictures		  Untitled Folder
Desktop				 Public			  Untitled.ipynb
Documents			 R			  Videos
Downloads			 shell-novice-data.zip	  wk2.ipynb
examples.desktop		 shell-novice-data.zip.1
exercise-01.ipynb		 temp
```

Renaming and copying files.

```python
!touch draft.txt
```

<!-- grade_id: 129, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir()
```

```python
!mv draft.txt quotes.txt
```

<!-- grade_id: 130, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "draft.txt" not in os.listdir()
```

```python
!ls
```

Output:
```
anaconda3			 Music			  temp
Anaconda3-4.1.1-Linux-x86_64.sh  numbers.txt		  Templates
data-shell			 Pictures		  Untitled1.ipynb
Desktop				 Public			  Untitled Folder
Documents			 quotes.txt		  Untitled.ipynb
Downloads			 R			  Videos
examples.desktop		 shell-novice-data.zip	  wk2.ipynb
exercise-01.ipynb		 shell-novice-data.zip.1
```

```python
!cp quotes.txt quotations.txt
```

<!-- grade_id: 134, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "quotations.txt" in os.listdir()
```

```python

```